# Podstawy Sztucznej Inteligencji 2020/2021


Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

---

## Iloczyn skalarny

Mając wektory $n$-elementowe $\mathbf{a}$ i $\mathbf{b}$ napisać funkcję obliczającą ich iloczyn skalarny:

$$s =\mathbf{a}\cdot \mathbf{b} =  \sum_{i=1}^{n} a_i b_i $$ 

 - za pomocą iteracji po elementach
 - używając funkcjonalności `numpy`

```python
import numpy as np
```

<!-- grade_id: cell-06e5e3a8f5681458, points: 0, grade: false, solution: true -->
```python
def s1(a,b):
    # YOUR CODE HERE
    if len(a) != len(b):
        raise ValueError("Wektory muszą być tej samej długości !")
    # a = [1,2,3,4]
    # b = [2,3,4,5]
    table = []
    for i in range(len(a)):
        table.append(a[i]*b[i])
    return np.sum(table)
    
    #raise NotImplementedError()

def s2(a,b):
    # YOUR CODE HERE
    if len(a) != len(b):
        raise ValueError("Wektory muszą być tej samej długości !")
    return np.dot(a,b)
    #raise NotImplementedError()
    
def s3(a,b):
    if len(a) != len(b):
        raise ValueError("Wektory muszą być tej samej długości !")
    c = zip(a,b)
    iloczyn = 0
    for i in c:
        iloczyn += i[0] * i[1]
    return iloczyn
```

```python
    a = [1,2,3,4]
    b = [2,3,4,5]
    list(zip(a,b))
```

Output:
```
[(1, 2), (2, 3), (3, 4), (4, 5)]
```

<!-- grade_id: cell-96f566befa6dfa4f, points: 4, grade: true, solution: false -->
```python
assert s1([1,1,1,1],[1,1,1,1]) == 4 
assert s2([1,1,1,1],[1,1,1,1]) == 4
assert np.isscalar(s2([1,1,1,1],[1,1,1,1]))
assert np.isscalar(s1([1,1,1,1],[1,1,1,1]))
assert s2([7,7,7],[7,7,7]) == 147
assert s1([7,7,7],[7,7,7]) == 147
```

## Mnożenie macierzy przez wektor:

Mając wektor $n$-elementowy $\mathbf{x}$ oraz macierz $\mathbf{A}$ o wymiarach $m\times n$ napisać funkcję obliczającą ich iloczyn:

$$\mathbf{y}= \mathbf{A}\mathbf{x}.$$

Każdy element wektora $\mathbf{y}$ jest dany przez:

$$y_i = \sum_{j=1}^{n} a_{ij} x_j $$ 

 - za pomocą iteracji po elementach (podwójna pętla)
 - korzystając z faktu, że
 każdy element wektora $y_i$ jest iloczynem skalarnym $i$-tego rzędu macierzy $\mathbf A$ oraz wektora $\mathbf{x}$ (pojedyncza pętla)
 - używając funkcji:  `np.dot` lub `np.tensordot` (bez pętli)


<!-- grade_id: cell-fb522d133295441a, points: 0, grade: false, solution: true -->
```python
def y1(A,x):
    # YOUR CODE HERE
    if A.shape[1] != x.shape[0]:
        raise ValueError
    table = []
    for n in A:
        table.append(sum(i*j for i,j in zip(n,x)))
    return table
            
    
    #raise NotImplementedError()

def y2(A,x):
    # YOUR CODE HERE
    table = []
    for n in A:
        table.append(np.dot(n,x))
    return table
    #raise NotImplementedError()
    
def y3(A,x):
    # YOUR CODE HERE
    return A.dot(x)
    #return np.dot(A,x)
    #raise NotImplementedError()
```

<!-- grade_id: cell-eff8a1c20f8fb28f, points: 6, grade: true, solution: false -->
```python
A = np.arange(12).reshape(3,4)
x = np.arange(4)
np.testing.assert_almost_equal(y1(A,x),np.array([14, 38, 62]))
np.testing.assert_almost_equal(y2(A,x),np.array([14, 38, 62]))
np.testing.assert_almost_equal(y3(A,x),np.array([14, 38, 62]))

```

<!-- grade_id: cell-982fed0f9cd60244, points: 1, grade: true, solution: false -->
```python
try:
    y1(np.arange(9).reshape(3,3),np.array([1,1]))
except ValueError:
    pass
else:
    raise AssertionError("Second dimension of matrix should be the same as dimension of vector!.")
```

## Mnożenie macierzy

Mamy dwie macierze $\mathbf{A}_{m\times n}$ i $\mathbf{B}_{n\times k}$. Napisać funkcję obliczającą ich iloczyn:

$$\mathbf{C}= \mathbf{A}\mathbf{B}$$.

Każdy element macierzy  $\mathbf{C}$ jest dany przez:

$$c_{ij} = \sum_{k=1}^{n} a_{ik} b_{kj} $$ 

 - za pomocą iteracji po elementach (potrójna pętla)
 - korzystając z faktu, że każdy element macierzy $c_{ij}$ jest iloczynem skalarnym $i$-tego rzędu macierzy $\mathbf A$ oraz  $j$-tego rzędu macierzy $\mathbf B$ (podwójna pętla)
 - używając funkcji:  `np.dot` lub `np.tensordot` (bez pętli)
 

<!-- grade_id: cell-bc5bc5d32ae48f5f, points: 0, grade: false, solution: true -->
```python
def C1(A,B):
    # YOUR CODE HERE
    if A.shape[1]!=B.shape[0]:
        raise ValueError
    table=[]
    for n in A:
        table.append(sum([i*j for i,j in zip (n,B)]))
    return np.array(table)
    #raise NotImplementedError()
    
def C2(A,B):
    # YOUR CODE HERE
    return np.array([np.dot(row,B) for row in A])
    #raise NotImplementedError()
    
def C3(A,B):
    # YOUR CODE HERE
    return np.dot(A,B)
    #raise NotImplementedError()
```

<!-- grade_id: cell-fdb6f13441a02162, points: 6, grade: true, solution: false -->
```python
m = np.arange(9).reshape(3,3)
assert np.prod(C1(np.eye(3),m) == m)
assert np.prod(C1(m,m) == C3(m,m))
assert np.prod(C2(m,m) == C3(m,m))


A,B = np.ones((4,5)),np.arange(10).reshape(5,2)
assert C1(A,B).shape == (4,2)
assert C2(A,B).shape == (4,2)
assert C3(A,B).shape == (4,2)

np.testing.assert_allclose(C1(A,B), A.dot(B))
np.testing.assert_allclose(C2(A,B), A.dot(B))
np.testing.assert_allclose(C3(A,B), A.dot(B))


```

<!-- grade_id: cell-f8df427efe4ce6de, points: 1, grade: true, solution: false -->
```python
try:
    C1(np.arange(9).reshape(3,3),np.ones((2,3)))
except ValueError:
    pass
else:
    raise AssertionError("Second dimension of matrix should be the same as dimension of vector!.")
```

## Ślad macierzy

Śladem macierzy nazywamy sumę elementów na przekątnej: 

$$ Tr(A) = \sum_{i=1}^{n} a_{ii}$$


Obliczyć ślad macierzy $n\times n$.



<!-- grade_id: cell-6cfbc4d0c3f5e8a4, points: 0, grade: false, solution: true -->
```python
def Tr(A):
    # YOUR CODE HERE
    if A.shape[0] != A.shape[1]:
        raise ValueError("Macierz musi być kwadratowa")
    return np.trace(A)
    #raise NotImplementedError()
```

<!-- grade_id: cell-6fcc4a8c88b8f364, points: 2, grade: true, solution: false -->
```python
assert Tr(np.diag([1,2,3,4])) == 10
```

<!-- grade_id: cell-3a491937c81d5cfd, points: 1, grade: true, solution: false -->
```python
try:
    Tr(np.ones((2,3)))
except ValueError:
    pass
else:
    raise AssertionError("The matrix should be square matrix!.")
```

## Wyznacznik macierzy

Obliczyć wyznacznik macierzy korzystając z rozwinięcia [Laplace'a](https://pl.wikipedia.org/wiki/Rozwini%C4%99cie_Laplace%E2%80%99a)


```python
import numpy as np
A = np.random.randint(4,size=(4,4)).astype(np.float)
x = np.random.randint(4,size=4)
```

<!-- grade_id: cell-a6b07fe35339bafe, points: 0, grade: false, solution: true -->
```python
def Le(A):
    # YOUR CODE HERE
    if A.shape[0] != A.shape[1]:
        raise ValueError("Wektory muszą być tej samej długości !")
    A= list(A)
    n = len(A)
    kopiaA = A.copy()
    for x in range(n):
        for i in range(x+1,n):
            if kopiaA[x][x] == 0:
                kopiaA[x][x] == 1.0e-18
            skala = kopiaA[i][x] / kopiaA[x][x] 

            for j in range(n):
                kopiaA[i][j] = kopiaA[i][j] - skala * kopiaA[x][j]
    wyznacznik = 1.0
    for i in range(n):
        wyznacznik *= kopiaA[i][i]
    return wyznacznik
    #raise NotImplementedError()
```

<!-- grade_id: cell-dac69a3e1b14c456, points: 10, grade: true, solution: false -->
```python
A = np.array([[ 1.,  3.,  0.,  0.],[ 3.,  2.,  3.,  2.],[ 1.,  2.,  2.,  2.],[ 1.,  0.,  0.,  3.]])
np.testing.assert_allclose( Le(A), np.linalg.det(A) )
```

```python

```